# BLU04 - Exercises Notebook

```python
import chardet
import hashlib # for grading purposes
import math
import numpy as np
import pandas as pd

from utils import f1, f2
```

## Q1: Use a shell command to count the number of lines in a file

<!-- grade_id: cell-498ef9d6b9a7ce6d, points: 0, grade: false, solution: true -->
```python
# Count the number of lines in file data/exercises/lusiadas.txt
# ! ...
# YOUR CODE HERE
raise NotImplementedError()

# Assign the number of lines to variable num_lines
# num_lines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a543516b4e915669, points: 2, grade: true, solution: false -->
```python
expected_hash = 'c715784eb51b33431fd4d41bac4467d6f6e86c9ee0e8db067cb3470da0f0c280'
assert hashlib.sha256(str(num_lines).encode()).hexdigest() == expected_hash
```

## Q2: Use a shell command to print the last three lines of a file

<!-- grade_id: cell-05a3d2245d57305d, points: 0, grade: false, solution: true -->
```python
# Print the last three lines of file data/exercises/lusiadas.txt
# ! ...
# YOUR CODE HERE
raise NotImplementedError()

# Assign to variable last_word_last_line the last word of the last line
# last_word_last_line = ...
# YOUR CODE HERE
raise NotImplementedError()

# Assign to variable last_word_second_last_line the last word of the second last line
# last_word_second_last_line = ...
# YOUR CODE HERE
raise NotImplementedError()

# Assign to variable last_word_third_last_line the last word of the third last line
# last_word_third_last_line = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e69c6137e73d90c, points: 2, grade: true, solution: false -->
```python
expected_hash = '716234e9338b59b6c3cb775c61da6ebfcc7f81f776497cd85bf5aaabf20f7991'
assert hashlib.sha256(last_word_last_line.encode()).hexdigest() == expected_hash

expected_hash = '79a1fbdd7c590c1206c7bceb7e72fab3af93c7bf846c25d0ea8b175657c5981b'
assert hashlib.sha256(last_word_second_last_line.encode()).hexdigest() == expected_hash

expected_hash = '9c75e4d30c06117cef345084d394cb2c3df8d64341921413a137087f5007ac87'
assert hashlib.sha256(last_word_third_last_line.encode()).hexdigest() == expected_hash
```

## Q3: Use the timeit magic to determine the execution time of functions

In this question, you should use the timeit line magic to determine the execution time of two functions: f1 and f2.
These functions were already imported for you in the first code cell of the Notebook.

<!-- grade_id: cell-f579359883e1007d, points: 0, grade: false, solution: true -->
```python
# Use the timeit line magic to determine the execution time of function f1()
# % ...
# YOUR CODE HERE
raise NotImplementedError()

# Use the timeit line magic to determine the execution time of function f2()
# % ...
# YOUR CODE HERE
raise NotImplementedError()

# Assign to variable fastest_function the name ('f1' or 'f2') of the function that had a lower execution time
# fastest_function = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9ff579d413c28ec, points: 2, grade: true, solution: false -->
```python
expected_hash = 'e4ab4e3b1493d5a997b4e51cdefbaa10570ef3ea9432bd72e7b6a89654ceb7f6'
assert hashlib.sha256(fastest_function.encode()).hexdigest() == expected_hash
```

## Q4: Read a delimiter separated file

Read file **data/exercises/euribor_interest_rates.csv** into a pandas DataFrame.

First, you should preview the file using a shell command in order to find out the used delimiter, and other properties of this file.

Then, you should use function read_csv to read the data into a DataFrame. The resulting DataFrame should have the first column as index.

<!-- grade_id: cell-942e0590467e20d9, points: 0, grade: false, solution: true -->
```python
# Use a shell command to preview the data
# ! ...
# YOUR CODE HERE
raise NotImplementedError()

# Use function read_csv to read the data into a DataFrame
# df4 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c275c9acaa9c74d, points: 2, grade: true, solution: false -->
```python
assert df4.loc[2002, 'Euribor 3 months'] == 2.87
assert set(df4.columns) == {'Euribor 3 months', 'Euribor 6 months', 'Euribor 12 months'}
assert len(df4) == 19

expected_hash = 'b68ca0811f132f73edc68e9d3bebb288ef036c1fef8aabe6d2c63a2b6bfa859c'
assert hashlib.sha256(df4.index.name.encode()).hexdigest() == expected_hash
```

## Q5: Read a csv file

Read file **data/exercises/portugal_urban_waste_per_inhabitant.csv** using function read_csv. Pay attention to the following:
* you might find some trouble when reading the file, at first, then just ignore the problems ;)
* use the first column as index
* there is a value in the file that should be interpreted as NaN, make sure you select the right one when reading the file!

<!-- grade_id: cell-2771b707556c08d2, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_urban_waste_per_inhabitant.csv with read_csv
# df5 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-601b354802964776, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df5.loc[1994, 'Urban waste collection per inhabitant (kg/inhabitant)'])
assert df5.loc[2011, 'Selective urban waste collection per inhabitant (kg/inhabitant)'] == 71.4

mean_selective_waste = df5['Selective urban waste collection per inhabitant (kg/inhabitant)'].mean()
assert math.isclose(37.0458, mean_selective_waste, rel_tol=1e-3)

expected_hash = 'b68ca0811f132f73edc68e9d3bebb288ef036c1fef8aabe6d2c63a2b6bfa859c'
assert hashlib.sha256(df5.index.name.encode()).hexdigest() == expected_hash
```

## Q6: Read a JSON file

Read file **data/exercises/portugal_production_of_electricity_gwh.json**.

<!-- grade_id: cell-5c2125479f7e50e5, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_production_of_electricity_gwh.json with read_json
# df6 = read_json(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3f085fa2cdad9319, points: 2, grade: true, solution: false -->
```python
assert len(df6) == 22
assert df6.loc[3, 'Hydropower > 10MW'] == 12488
assert set(df6.columns) == {'Biomass', 'Geothermal power', 'Hydropower < 10MW', 'Hydropower > 10MW', 
                           'Photovoltaic', 'Total','Total renewable sources','Windpower','Year'}
```

## Q7: Read an HTML file

Read file **data/exercises/portugal_cities.html** into a pandas DataFrame.
If you notice something weird in the characters, make sure you specify a suitable encoding.

<!-- grade_id: cell-fb39107093dd73fb, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_cities.html with read_json
# df7 = read_html(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ec15a8be42137a01, points: 2, grade: true, solution: false -->
```python
assert len(df7) == 151
assert df7.loc[1, 'City'] == 'Agualva-Cacém'
assert df7.loc[143, 'Population'] == 30100
```

## Q8: Read an Excel file

Read file **data/exercises/portugal_gas_emissions_per_year.xlsx** using function read_excel. Pay attention to the following:

* use column 'Years' as index
* make sure you discard the rows and columns with no data

<!-- grade_id: cell-017a7c31b0ddc38e, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/portugal_gas_emissions_per_year.xlsx with read_excel
# df8 = read_excel(...)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-91d5f07e40585680, points: 2, grade: true, solution: false -->
```python
assert math.isclose(209909.9, df8.loc[2006, 'Non-methane volatile organic compounds'], rel_tol=1e-3)
assert math.isclose(3562538.3, df8.sum(axis=1)[2012], rel_tol=1e-3)

mean_methane = df8['Methane'].mean()
assert math.isclose(492018.5, mean_methane, rel_tol=1e-3)
```

## Q9: Find the encoding of a file

Find the encoding used in file **data/exercises/mystery_cities.csv**, using the method that was shown in the Learning Units.

Then, read the data into a DataFrame, using the read_csv method and find the date in the third row, second column of the DataFrame.

<!-- grade_id: cell-bd6944fd63bb38d8, points: 0, grade: false, solution: true -->
```python
# Find the encoding of file data/exercises/mystery_cities.csv
# encoding = ...
# YOUR CODE HERE
raise NotImplementedError()

# Read the file into a DataFrame
# df9 = ...
# YOUR CODE HERE
raise NotImplementedError()

# Find the date in the third row, second column of the DataFrame
# date_found = df9.iloc[...]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-687fb83c97e03355, points: 2, grade: true, solution: false -->
```python
expected_hash = '8e4652c10a930984607dce82f395f3822bae1d4bedabca344105daf4b4520d53'
assert hashlib.sha256(encoding.encode()).hexdigest() == expected_hash

assert isinstance(df9, pd.DataFrame)

assert date_found == '1949-10-01'
```

## Q10: Clean data after reading it

Read file **data/exercises/world_percentage_of_literacy.tsv** and then clean it.
Note that:

* the missing values should be represented by NaNs
* the final DataFrame should have all the columns with datatype float, except the 'Country' column
* the index should be incremental starting from 0 (i.e, you don't need to read any column as the index)

In the end, answer the question: which country has the highest literacy rate gap between women and men (with higher literacy in men than in women)?

<!-- grade_id: cell-5106c25a705296e8, points: 0, grade: false, solution: true -->
```python
# Read file data/exercises/world_percentage_of_literacy.tsv and clean it!
# df10 = pd.read_csv(...)
# YOUR CODE HERE
raise NotImplementedError()

# Answer the literacy rate gender gap question
# hint: use idxmin()
# country_with_higher_literacy_rate_gender_gap = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f60a22c3465d5b83, points: 2, grade: true, solution: false -->
```python
assert np.isnan(df10.loc[2, 'Non-UNESCO literacy rate'])
assert df10.loc[23, 'Female literacy'] == 88.9
assert df10.dtypes['Literacy rate (all)'] == np.float
assert df10.dtypes['Male literacy'] == np.float
assert df10.dtypes['Female literacy'] == np.float
assert df10.dtypes['Gender difference'] == np.float
assert df10.dtypes['Non-UNESCO literacy rate'] == np.float

expected_hash = '5c2e4745c608d719cab6de3f2c11f3d972e212334f794dea783b7dc16c7b31f2'
assert hashlib.sha256(country_with_higher_literacy_rate_gender_gap.encode()).hexdigest() == expected_hash
```

```python

```